# SLU1 - Pandas 101: Exercise notebook

```python
# First things first, import:
```

```python
import pandas as pd
import numpy as np
from hashlib import sha256
import json
```

- In this notebook the following is tested:

    - Pandas Series
    - Pandas DataFrames
    - Printing the columns
    - Load a dataset
    - Preview a dataframe
    - Make use of info and describe

## Exercise 1: Series

In this first exercise the goal is to get used to creating series.

#### 1.1) Create a series with the name of the Avengers: 

<!-- grade_id: cell-d0be99db1c88ce43, points: 0, grade: false, solution: true -->
```python
# Create a list with the name of the avengers provided below then
# Create a series with that list, using the order provided.
# Let pandas generate the index on its own. 
# Don't forget to delete the raise error

# names: Black Widow, Iron Man, Hulk, Captain America, Black Panther.

# avengers_list = ...

# Create a series with members of the Avengers team, and call it avengers_series. 
# avengers_series = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
avengers_series
```

<!-- grade_id: cell-0344c841b5dac488, points: 2, grade: true, solution: false -->
```python
assert(isinstance(avengers_list, list)) ,"The data structure used it's not the correct one"
assert len(avengers_series.index ) == 5 ,"The lenght of the list is not correct"
assert ((avengers_series.values[3] == 'Captain America') & (avengers_series.values[-1] =='Black Panther')), "The order of the list is not correct"
assert (isinstance(avengers_series, pd.Series)), "avengers_series is not a pandas Series"
```

#### 1.2) Create a series for the countries of each super hero club (using the avengers_list as index):

<!-- grade_id: cell-2d2283c13fa0600e, points: 0, grade: false, solution: true -->
```python
# Data to create the pandas Series
country_list = ['Russia', 'USA', 'USA', 'USA', 'Wakanda']

# Create a series which has the avengers_list (previously created) as index, and the respective country as values. 
# country_series = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
country_series
```

<!-- grade_id: cell-096e084aeb31ba69, points: 2, grade: true, solution: false -->
```python
assert (isinstance(country_series, pd.Series)), "country_series is not a list"
assert(len(country_series==5)), "The lenght of country_series is not correct"
assert(country_series[-3] == 'USA'), 'The order of country_series is not correct'
assert(country_series.index[3] == 'Captain America'), 'The order of the index is not correct'
```

#### 1.3) Get an numpy array with the the index of the given country_series.

<!-- grade_id: cell-5882c53ea453ac16, points: 0, grade: false, solution: true -->
```python
# Get an array with the indexes of country_series.
# index_series = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
index_series
```

<!-- grade_id: cell-c733f3352254f868, points: 1, grade: true, solution: false -->
```python
assert (isinstance(index_series, np.ndarray)), "index_series is not in the format of a numpy array"
assert(len(index_series==5)), 'The lenght of index_series is not correct'
assert(index_series[2] == 'Hulk'), 'The order of index_series is not correct'
```

## Exercise 2: DataFrames

In this exercise the goal is to create a simple DataFrame.

#### 2.1) Create a DataFrame with avenger, country, age columns

<!-- grade_id: cell-b5c9767e3556cd9b, points: 0, grade: false, solution: true -->
```python
# Create a dataframe with 3 columns: avenger_list, country_list, age_list. 
# Let pandas create the index itself
# Make sure the columns are in the right order 
# avengers_df must be the name of the final dataframe.

# Remember that 2 of the necessary lists for the dataframe have been created already
## avengers_list
## countries_list

age_list = [34, 45, 37, 94, 20]

# build the dataframe, using a dictionary to tell it what should be the column names
# remember to give the columns the right names, or it won't pass the asserts
# avengers_df_dictionary = ...
# avengers_df = ...


# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# The expected output is a dataframe with columns named
# "avenger", "country" and "age", indexed from 0 to 4.
avengers_df
```

<!-- grade_id: cell-a4ab78569f17b22b, points: 2, grade: true, solution: false -->
```python
assert(isinstance(avengers_df, pd.DataFrame)), 'avengers_df is not a DataFrame'
assert(isinstance(avengers_df_dictionary,dict)), 'avengers_df_dictionary is not a dictionary'
assert(avengers_df['country'].tolist()==country_list), 'The lists used to create the dataframe are not associated to the correct columns'
assert(avengers_df['avenger'].tolist()==avengers_list), 'The lists used to create the dataframe are not associated to the correct columns'
assert(avengers_df['age'].tolist()==age_list), 'The lists used to create the dataframe are not associated to the correct columns'
assert(avengers_df.shape == (5, 3)), 'The size of the dataframe is not correct'
assert(avengers_df.iloc[3, 1] == 'USA'), 'The organization of the data appears not to be correct'
assert(avengers_df.columns[2] == 'age'), 'The organization of the data appears not to be correct'
```

#### 2.2) Create a DataFrame with the country/age columns, but with avenger as row indexes

<!-- grade_id: cell-0c74e28305cd2c4a, points: 0, grade: false, solution: true -->
```python
# Create a dataframe with columns country and age (in that order), but use avengers_list as index
# avengers_df_ind must be the name of the final dataframe.

# All the necessary values are already stored in their respective variables 

# Use the previous lists to create Series for country and age
# Hint: use the avengers_list as index
# country_series = ...
# age_series = ...

# Finally create the DataFrame, using a dictionary (remember that this dictionary should include the series) 
# avengers_df_ind_dictionary = ...
# avengers_df_ind = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# run this cell to see the results 
# The expected output is a dataframe with columns named "country", "age", indexed with the correct avengers names.
avengers_df_ind
```

<!-- grade_id: cell-0193acd630b4eb40, points: 2, grade: true, solution: false -->
```python
assert(isinstance(avengers_df_ind_dictionary,dict)), 'avengers_df_ind_dictionary is not a dictionary'
assert([n for n in avengers_df_ind_dictionary.keys()][0]=='country'),'the order of the columns is not correct'
assert(avengers_df_ind['country'].tolist()==country_list), 'The lists used to create the dataframe are not associated to the correct columns'
assert(isinstance(avengers_df_ind, pd.DataFrame)), 'avengers_df_ind is not a DataFrame'
assert(avengers_df_ind.shape == (5, 2)), 'The size of the DataFrame is not correct'
assert(avengers_df_ind.index[3] == 'Captain America'), 'The order of the index is not correct'
assert(avengers_df_ind.columns[1] == 'age'), 'The order of the columns is not correct'
backup_avengers_list = avengers_list.copy()
```

## Exercise 3: Load the Digimon dataset


In this exercise you will load the Digimon data set.

You will then preview it and retrieve some information.

#### 3.1) Load the data set

<!-- grade_id: cell-6f209833953ab9cf, points: 0, grade: false, solution: true -->
```python
# Load the Digimon data set
# it is located at data/DigiDB_digimonlist.csv
# digimon_df = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# print head it to get an idea of what you've just loaded 
digimon_df.head()
```

<!-- grade_id: cell-a08b45d8b2814c8f, points: 1, grade: true, solution: false -->
```python
assert(isinstance(digimon_df, pd.DataFrame)),'digimon_df is not a DataFrame'
assert(digimon_df.shape == (249, 13)), 'the size of digimon df is not correct'
```

#### 3.2) Load the data set with an index defined

<!-- grade_id: cell-8e15e2392133ead5, points: 0, grade: false, solution: true -->
```python
# Load the digimon data set and set the index the "Number" column 
# (check the documentatio: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html)
# digimon_df_ind = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# print head it to get an idea of what you've just loaded 
digimon_df_ind.head()
```

<!-- grade_id: cell-4457e5154efddfdf, points: 1, grade: true, solution: false -->
```python
assert(isinstance(digimon_df_ind, pd.DataFrame)),'digimon_df is not a DataFrame'
assert(digimon_df_ind.shape == (249, 12)), 'the size of digimon df is not correct'
```

#### 3.3) Get general information about the dataframe

<!-- grade_id: cell-09bcf295f244ae5e, points: 0, grade: false, solution: true -->
```python
# NOT GRADED

# Print digimon_df_ind info
# reminder: jupyter just prints out the last output, 
# so you don't need to use a print statement
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

#### 3.4) Preview the top 7 entries

<!-- grade_id: cell-84111368d39ecdb5, points: 0, grade: false, solution: true -->
```python
# Print the top 7 entries of digimon_df_ind

# digimon_df_ind_top = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
digimon_df_ind_top
```

<!-- grade_id: cell-5354f5e7e9461336, points: 1, grade: true, solution: false -->
```python
assert (digimon_df_ind_top.shape == (7,12)), 'The output of the digimon_df_ind_top is not correct'
assert (digimon_df_ind_top.iloc[-1,0] == 'Tanemon'), 'The organization of the data is not correct'
```

#### 3.5) Preview the bottom 3 entries

<!-- grade_id: cell-afd010e3e4e4c1f0, points: 0, grade: false, solution: true -->
```python
# Print the bottom 3 entries of digimon_df_ind

# digimon_df_ind_bottom = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0e1de9cf480b483d, points: 1, grade: true, solution: false -->
```python
assert (digimon_df_ind_bottom.shape == (3,12)), 'The output of the digimon_df_ind_top is not correct'
assert (digimon_df_ind_bottom.iloc[-1,0] == 'Rosemon BM'), 'The organization of the data is not correct'
```

###### 3.6) Use a method to determine all the types os data in digimon_df_ind

<!-- grade_id: cell-e9e62f9d35948d34, points: 0, grade: false, solution: true -->
```python
# Print the type of data in digimon_df_ind
#digimon_df_ind_types = ...

# Manually insert the answers to the following variables:

# number_of_cols = ...
# number_of_obj_cols = ...
# number_of_int_cols = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-efab0f02db3e97d7, points: 2, grade: true, solution: false -->
```python
assert (number_of_cols==12), 'Total number of columns is not correct'
assert (number_of_obj_cols==4), 'Number of object type columns is not correct'
assert (number_of_int_cols==8), 'Number of integer columns is not correct'
```

#### 3.7) Print the number of rows and the number of columns

<!-- grade_id: cell-e41c90b0d616df98, points: 0, grade: false, solution: true -->
```python
# Get the nr_rows and nr_cols from digimon_df_ind, using a method

# nr_rows = ...
# nr_cols = ...

#YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b8e256291698a0b, points: 1, grade: true, solution: false -->
```python
assert (nr_rows == 249), 'Number of rows in digimon_df_ind is not correct'
assert (nr_cols == 12), 'Number of columns in digimon_df_ind is not correct'
```

#### 3.8) Get the array of column names

<!-- grade_id: cell-4fd5c93627335cd6, points: 0, grade: false, solution: true -->
```python
# Create an array with the column names of digimon_df_ind_columns and print it.
# there is, as always, a method for this 
# digimon_df_ind_columns = ...

#YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(digimon_df_ind_columns)
```

The expected output is an array with 'Digimon', 'Stage', 'Type', 'Attribute', 'Memory', 'Equip Slots',
       'Lv 50 HP', 'Lv50 SP', 'Lv50 Atk', 'Lv50 Def', 'Lv50 Int', 'Lv50 Spd'.

<!-- grade_id: cell-4dce9e138896815d, points: 2, grade: true, solution: false -->
```python
assert(isinstance(digimon_df_ind_columns,np.ndarray)), 'digimon_df_ind_columns is not in the format of a numpy array'
assert (digimon_df_ind_columns[1] == 'Stage' ), 'the order of digimon_df_ind_columns is not correct'
assert (len(digimon_df_ind_columns) == 12), 'The lenght of digimon_df_ind_columns is not correct'
```

#### 3.9) Making use of describe() what is the standard deviation value of the Memory?

```python
# Print digimon_df_ind describe

digimon_df_ind.describe()
```

<!-- grade_id: cell-3681531b79d28fd3, points: 0, grade: false, solution: true -->
```python
# Manually input the standard deviation value of the Memory column
# yep, that simple. No tricks here :) 

#memory_std = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1bf51b9a2de225bc, points: 2, grade: true, solution: false -->
```python
assert 6 < memory_std < 7, 'The value for standard deviation of the column memory is wrong'
```