## SLU17 - Ethics & Fairness - Exercise notebook

```python
import warnings
warnings.filterwarnings('ignore')
import hashlib

import matplotlib.pyplot as plt
plt.style.use('seaborn-dark')
%matplotlib inline

import pandas as pd
from sklearn.metrics import confusion_matrix
import numpy as np

from utils.utils import make_data
```

# Criminal justice bias

Exercise adaptated from the book [Fairness and Machine Learning by Solon Barocas, Moritz Hardt, and Arvind Narayanan](https://fairmlbook.org/pdf/fairmlbook.pdf).

Based on the ProPublica's article [Machine Bias](https://www.propublica.org/article/machine-bias-risk-assessments-in-criminal-sentencing) about a proprietary risk score, called COMPAS, used in the US.

This is the problem setting:

> Risk assessment is an important component of the criminal justice system. In the United States, judges set bail and decide pre-trial detention based on their assessment of the risk that a released defendant would fail to appear at trial or cause harm to the public.

These scores are intended to assess the risk that a defendant will re-offend, a task often called **recidivism prediction**.

We’ll use data obtained and released by ProPublica.

```python
data = make_data()
```

## Exercise 1 - Score distribution

### Score distribution

Plot the histogram of the scores, `decile_score`.

Suggestion: Use `Series.hist(histtype='step')`.

<!-- grade_id: cell-1214438aeb6fb1fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Score distribution by race

Plot the histogram of the scores for black (`race == "African-American"`) and white defendants (`race == "Caucasian"`).

#### White Defendants’ Risk Scores

<!-- grade_id: cell-c0db9a1ce9d279ff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Black Defendants’ Risk Scores

<!-- grade_id: cell-e73a94d85986910b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Interpretation

Based on these plots, what would you conclude from the distributions. Uncomment the correct answer.

<!-- grade_id: cell-5a21d6a873bee7f1, points: 0, grade: false, solution: true -->
```python
# hypothesis_1 = 'The distribution of the scores are similar for both populations.'
# hypothesis_1 = 'Scores for white defendants are skewed toward lower-risk categories.'
# hypothesis_1 = 'Scores for black defendants are skewed toward lower-risk categories.'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c7d5b413b2522590, points: 5, grade: true, solution: false -->
```python
expected_hash_1 = '1d0a31bdfb0e6cd81f854bad389789d7fd0022bdf92534fc6fbd0c80bc40e817'
assert hashlib.sha256(str(hypothesis_1).encode('utf-8')).hexdigest() == expected_hash_1
```

## Exercise 2 - Error rates

### Overall

Plot the histogram of scores received by the positive class (recidivists). Use the column `two_year_recid`.

<!-- grade_id: cell-0d47432f9e15cc0b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Overall, the risk score doesn't appear to be particularly good at separating recidivists.

(The resulting histogram should resemble a uniform distribution.)

### By race

Plot the same histogram for black (`race == "African-American"`) and white defendants (`race == "Caucasian"`).

#### White Defendants

<!-- grade_id: cell-c75dac65af59a0ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Black Defendants

<!-- grade_id: cell-c58fb419278988c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Interpretation

Based on these results, uncomment the correct answer.

(Remember, these histograms report to **actual recidivists**.)  

<!-- grade_id: cell-60a69c870b9bd3ca, points: 0, grade: false, solution: true -->
```python
# hypothesis_2 = 'The distribution of recidivists scores is similar in both groups.'
# hypothesis_2 = 'Scores for white recidivists are skewed toward lower-risk categories.'
# hypothesis_2 = 'Scores for black recidivists are skewed toward lower-risk categories.'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ab17b0f1269533b3, points: 5, grade: true, solution: false -->
```python
expected_hash_2 = '8fb14ceb9ee8aefb715152afa576bb0f7af71bb7671062a8d94c1b813897b936'
assert hashlib.sha256(str(hypothesis_2).encode('utf-8')).hexdigest() == expected_hash_2
```

## Exercise 3 - When predictions fail differently

Defendants with `decile_score` higher than 3 are classified as high-risk of recidivism.

We use `scikit-learn` to compute the confusion matrix by race for us.

### False Positive Rate (FPR)

The false positive rate, also known as probability of false alarm, is given by:

$$FPR = \frac{FP}{FP + TN} = \frac{FP}{N}$$

Where $FP$ is the number of false positives, $TN$ is the number of true negatives, and $N$ the total number of negatives.

Below, we compute the global FPR of the `decile_score`.

```python
y_true = data['two_year_recid']
y_pred = data['decile_score'] > 3

tn, fp, fn, tp = confusion_matrix(y_true, y_pred).ravel()

fpr = fp / (fp + tn)
```

#### White Defendants

Compute the FPR for white defendants.

<!-- grade_id: cell-2c45aed7cea89b5d, points: 0, grade: false, solution: true -->
```python
# fpr_w = ...
# YOUR CODE HERE
raise NotImplementedError()
```

```python
fpr_w
```

<!-- grade_id: cell-9f195b34ea5e7d27, points: 5, grade: true, solution: false -->
```python
assert np.isclose(fpr_w, 0.33, atol=0.01)
```

#### Black Defendants

Compute the FPR for black defendants.

<!-- grade_id: cell-f311ec4dff1d2d26, points: 0, grade: false, solution: true -->
```python
# fpr_b = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-35975da0f9218ac3, points: 5, grade: true, solution: false -->
```python
assert np.isclose(fpr_b, 0.54, atol=0.01)
```

```python

```