```python
import hashlib # for grading purposes
```

# Main Goal

The main idea of the workflow unit is for you to understand how you should build a strategy for designing a Recommender System.
For this purpose we recommend that you:

1. Review the workflow notebook carefully and make sure you are able to explain to a 5-year-old child (or a [rubber duck](https://en.wikipedia.org/wiki/Rubber_duck_debugging)) the whole pipeline. Even better, try to get some ideas on how to improve it!

2. Choose one of these [datasets](https://github.com/caserec/Datasets-for-Recommender-Systems) and follow the steps on the Simplified Pipeline notebook.

3. Design and implement the whole RS pipeline for another dataset from stratch - as will happen on the Hackathon.

When you feel ready for the Hackathon complete the next cell and submit it to the Portal.


<!-- grade_id: cell-e7c828874307647a, points: 0, grade: false, solution: true -->
```python
def ds_recommenders_oath_of_office():
    """The Recommender Systems' Data Scientist Oath of Office
    
    Returns:
        oath (str): should contain the following sentence:
                    'I do solemnly swear I have reviewed the Workflow notebook and I am ready for the Hackathon!'
    """
    # YOUR CODE HERE
    raise NotImplementedError()

oath = ds_recommenders_oath_of_office()
```

<!-- grade_id: cell-5ce403cc55ee4d49, points: 20, grade: true, solution: false -->
```python
expected_hash = '3829f34c0b9c1e4a4da31670486d219446ae6ad92b0a243d92ec41793611ed7f'
assert hashlib.sha256(oath.encode()).hexdigest() == expected_hash
```