# SLU10 - Metrics for Regression: Exercise Notebook

In this notebook, you will implement:
    - Mean Absolute Error (MAE)
    - Mean Squared Error (MSE)
    - Root Mean Squared Error (RMSE)
    - Coefficient of Determination (R²)
    - Adjusted R²
    - Scikitlearn metrics
    - Using metrics for k-fold cross validation


Start by loading the data we will use to fit a linear regression - hopefully you still have SLU07 in your memory - and fitting the LinearRegression estimator from scikitlearn:

```python
# Base imports
import math
import numpy as np
import pandas as pd
from sklearn.datasets import load_diabetes
from sklearn.linear_model import LinearRegression

data = load_diabetes()

x = pd.DataFrame(data['data'], columns=data['feature_names'])
y = pd.Series(data['target'])
x.head()
```

```python
np.random.seed(42)

x_diabetes = x.values
y_diabetes = y.values

lr = LinearRegression()
lr.fit(x_diabetes, y_diabetes)

y_hat_diabetes = lr.predict(x_diabetes)
betas_diabetes = pd.Series([lr.intercept_] + list(lr.coef_))
```

## 1 Metrics

We will start by covering the metrics we learned in the unit, in particular a set of related metrics:

- Mean Squared Error

$$MSE = \frac{1}{N} \sum_{n=1}^N (y_n - \hat{y}_n)^2$$


- Root Mean Squared Error

$$RMSE = \sqrt{MSE}$$


- Mean Absolute Error

$$MAE = \frac{1}{N} \sum_{n=1}^N \left| y_n - \hat{y}_n \right|$$

### 1.1 Mean Squared Error

Implement the mean squared error in the next function:

<!-- grade_id: cell-d7a35367fa5f0f10, points: 0, grade: false, solution: true -->
```python
def mean_squared_error(y, y_pred): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        
    Returns: 
        mse : float with Mean Squared Error Value

    """
    # 1) Compute the error.
    # error = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2) Compute the squared value of the errors for each sample
    # squared_error = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 3) Compute the mean squared value of the errors
    # mse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return mse
```

Check the outputs of your function match the results below:

<!-- grade_id: cell-8ea8afca85e2b4bd, points: 2, grade: true, solution: false -->
```python
mse = mean_squared_error(y_diabetes, y_hat_diabetes)
print('Mean Squared Error Diabetes dataset: {}'.format(mse))
assert math.isclose(2859.6903987680657, mse)
```

### 1.2 Root Mean Squared Error
Implement the root mean squared error in the function below:

<!-- grade_id: cell-116b13594e8c311c, points: 0, grade: false, solution: true -->
```python
def root_mean_squared_error(y, y_pred): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        
    Returns: 
        mse : float with the Root Mean Squared Error Value
    """
    # 1) Compute the mean squared error. 
    # mse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2) Compute the root square.
    # rmse = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return rmse
```

Check the outputs of your function match the results below:

<!-- grade_id: cell-774c0900c5f9193c, points: 2, grade: true, solution: false -->
```python
rmse = root_mean_squared_error(y_diabetes, y_hat_diabetes)
print('Root Mean Squared Error Diabetes dataset: {}'.format(rmse))
assert math.isclose(53.47607314274362, rmse)
```

### 1.3 Mean Absolute Error

Finally, implement the Mean Absolute Error in the function below. 

<!-- grade_id: cell-4959a7de3c3ad479, points: 0, grade: false, solution: true -->
```python
def mean_absolute_error(y, y_pred): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        
    Returns: 
        mae : float with Mean Absolute Error
    """
    # 1) Compute the error.
    # error = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2) Compute the absolute value of the errors for each sample
    # abs_error = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 3) Compute the mean of the absolute value of the errors
    # mae = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return mae
```

Check the outputs of your function match the results below:

<!-- grade_id: cell-41455e2d039c8354, points: 2, grade: true, solution: false -->
```python
mae = mean_absolute_error(y_diabetes, y_hat_diabetes)
print('Mean Absolute Error Diabetes dataset: {}'.format(mae))
assert math.isclose(43.277395083749866, mae)
```

Next we will focus on the Coefficient of Determination - $R^2$ - and its adjusted form. See the equations below:

- $R^2$ score 

$$R² = 1 - \frac{MSE(y, \hat{y})}{MSE(y, \bar{y})} 
= 1 - \frac{\frac{1}{N} \sum_{n=1}^N (y_n - \hat{y}_n)^2}{\frac{1}{N} \sum_{n=1}^N (y_n - \bar{y})^2}
= 1 - \frac{\sum_{n=1}^N (y_n - \hat{y}_n)^2}{\sum_{n=1}^N (y_n - \bar{y})^2}$$

where $$\bar{y} = \frac{1}{N} \sum_{n=1}^N y_n$$

- Adjusted $R^2$ score 

$$\bar{R}^2 = 1 - \frac{N - 1}{N - K - 1} (1 - R^2)$$

where $N$ is the number of observations in the dataset used for training the model (i.e. number of rows of the pandas dataframe) and $K$ is the number of features used by your model (i.e. number of columns of the pandas dataframe)


### 1.4 R² score

Start by implementing the $R^2$ score in the function below:

<!-- grade_id: cell-51f11aef440a31c5, points: 0, grade: false, solution: true -->
```python
def r_squared(y, y_pred): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        
    Returns: 
        r2 : float with R squared value
    """

    # 1) Compute labels mean.
    # y_mean = ...
    # YOUR CODE HERE
    raise NotImplementedError()

    # 2) Compute the mean squared error between the target and the predictions.
    # mse_pred = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 3) Compute the mean squared error between the target and its mean.
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 4) Finally, compute R²
    # r2 = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return r2
```

Check the outputs of your function match the results below:

<!-- grade_id: cell-7fb94c0cc204ea9b, points: 3, grade: true, solution: false -->
```python
r2 = r_squared(y_diabetes, y_hat_diabetes)
print('R² Diabetes dataset: {}'.format(r2))
assert math.isclose(0.5177494254132934, r2)
```

### 1.5 Adjusted R² score

Then implement the adjusted $R^2$ score in the function below:

<!-- grade_id: cell-b35baba302fa675b, points: 0, grade: false, solution: true -->
```python
def adjusted_r_squared(y, y_pred, K):
    """
    Args: 
        y : numpy.array with shape (num_samples,) - labels 
        y_pred : numpy.array with shape (num_samples,) - predictions
        K : integer - Number of features used in the model that computed y_hat.

    Returns: 
        r2_adj : float with adjusted R squared value
    """
    
    # 1) Compute R².
    # r2 = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2) Get number of samples 
    # N = ...
    # YOUR CODE HERE
    raise NotImplementedError()

    # 3) Adjust R²
    # r2_adj = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return r2_adj
```

Check the outputs of your function match the results below:

<!-- grade_id: cell-2f310b9310e8fc38, points: 3, grade: true, solution: false -->
```python
r2 = adjusted_r_squared(y_diabetes, y_hat_diabetes, x_diabetes.shape[1])
print('Adjusted R² Diabetes dataset: {}'.format(r2))
assert math.isclose(0.506560316954205, r2)
```

## 2 Scikit-Learn metrics

As you know, scikitlearn also already provides you with implementations of these metrics: 

- `sklearn.metrics.mean_absolute_error`
- `sklearn.metrics.mean_squared_error`
- `sklearn.metrics.r2_score`
- `sklearn.linear_model.LinearRegression.score` 

```python
# Import sklearn metrics
from sklearn import metrics as sklearn_metrics
```

#### 2.1 Root Mean Squared Error

Implement the root mean squared error functions below with scikitlearn:

<!-- grade_id: cell-91685e84daf075d0, points: 0, grade: false, solution: true -->
```python
def sklearn_root_mean_squared_error(y, y_pred): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        
    Returns: 
        rmse : float with Root Mean Squared Error
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Make sure your function passes the tests below:

<!-- grade_id: cell-bb0647e6f49b6b36, points: 2, grade: true, solution: false -->
```python
rmse = sklearn_root_mean_squared_error(y_diabetes, y_hat_diabetes)
print('Sklearn RMSE Diabetes dataset: {}'.format(rmse))
assert math.isclose(53.47607314274362, rmse)
```

#### 2.2  Adjusted R² score

Implement the adjusted R² score below using scikitlearn:

<!-- grade_id: cell-bc8ebf5c8c8c6e23, points: 0, grade: false, solution: true -->
```python
def sklearn_adjusted_r_squared(y, y_pred, K): 
    """
    Args: 
        y_pred : numpy.array with shape (num_samples,) - predictions
        y : numpy.array with shape (num_samples,) - labels 
        K : integer - Number of features used in the model that computed y_hat.

    Returns: 
        r2_adj : float with adjusted R squared value
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Make sure your function passes the tests below:

<!-- grade_id: cell-f53ac73f1b9e0e25, points: 3, grade: true, solution: false -->
```python
r2 = sklearn_adjusted_r_squared(y_diabetes, y_hat_diabetes, x_diabetes.shape[1])
print('Sklearn Adjusted R² Diabetes dataset: {}'.format(r2))
assert math.isclose(0.506560316954205, r2)
```

Finally, compare the sklearn-based metrics with your own for the diabetes dataset:

```python
MAE = mean_absolute_error(y_diabetes, y_hat_diabetes)
MSE = mean_squared_error(y_diabetes, y_hat_diabetes)
RMSE = root_mean_squared_error(y_diabetes, y_hat_diabetes)
R2 = r_squared(y_diabetes, y_hat_diabetes)
R2_adj = adjusted_r_squared(y_diabetes, y_hat_diabetes, x_diabetes.shape[1])

print('Metric for diabetes dataset with base implementation:')
print('Mean Absolute Error diabetes dataset: {}'.format(MAE))
print('Mean Squared Error diabetes dataset: {}'.format(MSE))
print('Root Mean Squared Error diabetes dataset: {}'.format(RMSE))
print('R² diabetes dataset: {}'.format(R2))
print('Adjusted R² diabetes dataset: {}'.format(R2_adj))
print('\n')

SK_MAE = sklearn_metrics.mean_absolute_error(y_diabetes, y_hat_diabetes)
SK_MSE = sklearn_metrics.mean_squared_error(y_diabetes, y_hat_diabetes)
SK_RMSE = sklearn_root_mean_squared_error(y_diabetes, y_hat_diabetes)
SK_R2 = sklearn_metrics.r2_score(y_diabetes, y_hat_diabetes)
SK_R2_adj = sklearn_adjusted_r_squared(y_diabetes, y_hat_diabetes, x_diabetes.shape[1])

print('Metric for diabetes dataset with scikitlearn:')
print('Mean Absolute Error diabetes dataset: {}'.format(SK_MAE))
print('Mean Squared Error diabetes dataset: {}'.format(SK_MSE))
print('Root Mean Squared Error diabetes dataset: {}'.format(SK_RMSE))
print('R² diabetes dataset: {}'.format(SK_R2))
print('Adjusted R² diabetes dataset: {}'.format(SK_R2_adj))

```

## 3 Using the Metrics

Now you'll use the metrics to fit and check performance of your LinearRegression and SGDRegressor, with the `cross_val_scores` method of scikitlearn. Implement the missing steps below:


<!-- grade_id: cell-3897cdeea2d9f3c1, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import cross_val_score
from sklearn import metrics
from sklearn import linear_model

def estimator_cross_fold(X, y, K, clf_choice='linear', scoring='neg_mean_squared_error'):
    """
    Args: 
        X : numpy.array with shape (num_samples, num_features) - sample data
        y : numpy.array with shape (num_samples,) - sample labels 
        K : integer - Number of iterations for k-fold
        clf_choice: choice of estimator 
        scoring : scoring function as per sklearn notation

    Returns: 
        clf: estimator trained with full data
        scores : scores for each fold
    """
    
    if clf_choice == 'linear':
        clf = linear_model.LinearRegression()
    elif clf_choice == 'sgd':
        clf = linear_model.SGDRegressor(max_iter=10000)
    else:
        print('Invalid estimator')
        return None
     
    # 1) Fit linear_model
    # YOUR CODE HERE
    raise NotImplementedError()

    # 2) Run k-fold cross validation
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return clf, scores
```

Let's run the k-fold cross validation for the several cases and get the average error:

<!-- grade_id: cell-5565c9de30569d21, points: 3, grade: true, solution: false -->
```python
np.random.seed(42)

clf_lr, nmse_lr = estimator_cross_fold(x_diabetes, y_diabetes.ravel(), 5, clf_choice='linear', scoring='neg_mean_squared_error')
assert math.isclose(-2993.0729432998864, nmse_lr.mean())

clf_sgd, nmse_sgd = estimator_cross_fold(x_diabetes, y_diabetes.ravel(), 5, clf_choice='sgd', scoring='neg_mean_squared_error')
assert math.isclose(-3024.1665383692743, nmse_sgd.mean())

clf_lr, r2_lr = estimator_cross_fold(x_diabetes, y_diabetes.ravel(), 5, clf_choice='linear', scoring='r2')
assert math.isclose(0.48231812211149394, r2_lr.mean())

clf_sgd, r2_sgd = estimator_cross_fold(x_diabetes, y_diabetes.ravel(), 5, clf_choice='sgd', scoring='r2')
assert math.isclose(0.4754722767620816, r2_sgd.mean())

```

```python
print('Cross val evaluation for diabetes dataset:')
print('NMSE with Linear Regression: {}'.format(nmse_lr.mean()))
print('NMSE with SGD: {}'.format(nmse_sgd.mean()))
print('R² Score with Linear Regression: {}'.format(r2_lr.mean()))
print('R² Score with SGD: {}'.format(r2_sgd.mean()))
```

For this particular case it seems that the linear regression generalises better than the SGD regressor. It's important to remind that the SGD regressor is at a slight disadvantage, because we didn't check the data distribution to understand if it has appropriate scaling. Remember that SGD will be sensitive to this, while linear regression won't. Feel free to replicate these exercises but applying min-max scaling beforehand and check the new results.