# SLU00 - LU Tutorial: Exercise notebook

The Exercise notebooks are the most delicate part of the Learning Unit. It is VERY important that these exercise notebooks are well-reviewed, understandable, and that students can get immediate feedback on whether or not their answers are right.

The Exercise notebooks look very different for the students and the instructors. For the instructors, they look like something like this:

<img width="500" src="https://i.imgur.com/kTxbSON.png"/>

While for students they look like this:

<img width="500" src="https://i.imgur.com/LtfvXgg.png"/>

## Basic Rules

### Make stuff read-only

Make as many cells read-only as possible. All markdown cells should be read-only. Many code cells that you don't want the students to be able to edit should be read-only. What the notebook should end up looking like is mostly read-only with sparsly distributed editable code cells in which the students implement focused exercises.

### Mirror the topics in the Learning Notebook

Just as in the Exercise notebooks, there should be a bulleted list that summarizes the topics that they will be tested on. In practice this will usually mirror the same bulleted list of topics that are covered in the Exercise notebook.

### NEVER use hidden cells for grading

This makes things a bit complicated to create graded cells that don't give away the answer but hidden cells make it impossible for the students to get immediate feedback on their answers which is a non-negotiable requirement.

### Make the exercises self-contained if possible

If a student cannot get the answer to an exercise in the middle of the notebook, it would be nice if they were able to continue on with the next ones. This may not always be possible but you should try to do so whenever possible.

### Make every exercise notebook sum to 20 points

If a student cannot get the answer to an exercise in the middle of the notebook, it would be nice if they were able to continue on with the next ones. This may not always be possible but you should try to do so whenever possible.

```python
# Place any important imports at the top of the notebook when possible
import pandas as pd
import hashlib
import json
```

In this notebook the following is tested:

- Basic list usage
- Dictionary definition
- Addition and subtraction

## Basic List Usage

The first example Exercise will focus around list usage. The concept of an Exercise will serve to define a section where there will be a set of questions in it. The Exercises will be numbered and the actual questions will also be numbered subsections. This allows students to be able to say things like: "In Exercise 3.2 I am having trouble understanding why it is wrong".

## Exercise 1: List Definition

In this first exercise the goal is to demonstrate a basic understanding of the list data structure and how to perform the most important and basic functions with it.

#### 1.1) Create a list of numbers

The format of a cell where the students are to write code that is graded is the following:

1. Some comments with general guidance.
1. A commented out declaration of the variable name they are to populate with the correct answer.
    - The correct answer.
    - Printed hashes of the correct answer or subsets of the correct answer. This part can be kind of tricky to figure out how how to do fairly and correctly. It can also be a bit difficult to subset the answer and check only sub-parts of it but you really should in order to avoid a situation where it is all-or-nothing for the student and they have no idea what part of the answer they could have gotten wrong.

I'm including a screenshot here so that anyone reading the student version of this notebook has an idea of what is going on behind the scenes. If you're looking at the instructor version of this notebook, the following image is probably redundant:

<img width="500px" src="https://i.imgur.com/nu91L9E.png"/>

<!-- grade_id: cell-6e62f4a661edbaf0, points: 0, grade: false, solution: true -->
```python
# Create a list of numbers in a particular order

# Create a list of integers 0-5 inclusive and assign it to the variable "numbers"
# numbers = ...

# YOUR CODE HERE
raise NotImplementedError()
```

Next we have an auto-graded cell where we will check answers. In order for the students to get immediate feedback on whether or not they have answered a question correctly, we will manually copy-paste the printed hashes above into variables so that we can check the student answers against them.

Don't worry, the print statements are removed when creating the student version of this notebook.

<!-- grade_id: cell-e1f0144e6b6e35ff, points: 10, grade: true, solution: false -->
```python
first_element_hash = 'e3b0c44298fc1c149afbf4c8996fb92427ae41e4649b934ca495991b7852b855'
numbers_hash = 'b0229c06acf4d5c98db175d4d054cd39e7fff51d8c37754b4a4b424f7967710d'

assert isinstance(numbers, list), "Should be of type list"
assert first_element_hash == hashlib.sha256(bytes(numbers[0])).hexdigest()
assert numbers_hash == hashlib.sha256(json.dumps(numbers).encode()).hexdigest()
```

1.1) Reverse sort the list of numbers

<!-- grade_id: cell-76c1342f063d35b5, points: 0, grade: false, solution: true -->
```python
# Create a list of numbers in a particular order

# Create a list of integers 0-5 inclusive, reverse the order of the list, and assign
# the result to a variable called numbers_reversed
# numbers_reversed = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-33470e71829ec84a, points: 10, grade: true, solution: false -->
```python
reversed_first_element_hash = '8855508aade16ec573d21e6a485dfd0a7624085c1a14b5ecdd6485de0c6839a4'
reversed_numbers_hash = 'b1df81229a625d18d32059097d1e89a7c7dea45dd3c75bdcfcd8fa33bcfcc1db'

assert isinstance(numbers, list), "Should be of type list"
assert reversed_first_element_hash == hashlib.sha256(bytes(numbers_reversed[0])).hexdigest()
assert reversed_numbers_hash == hashlib.sha256(json.dumps(numbers_reversed).encode()).hexdigest()
```