# SLU 05 - Covariance and Correlation: Exercise notebook

```python
import pandas as pd 
import numpy as np 
import math 
import utils 

from matplotlib import pyplot as plt 
from utils import get_house_prices_and_rooms, plot_house_prices_and_rooms

# this is for grading without showing the answers 
import hashlib
def hash_answer(answer): 
    answer=str(answer)
    return hashlib.sha256((answer).encode()).hexdigest()
```

In this notebook you will practice the following: 

    - Covariance 
    - Pearson correlation
    - Spearman correlation
    - Correlation matrix
    - Spurious correlations

----

# Exercise 1: implement covariance 
Here you will implment covariance, by completing the following function.

Don't worry too much about generalizing, this will be an extremely naïve implementation, just to get your hands dirty!

Here is a quick reminder of the formula: 
$$ cov = \frac{\sum{(X - X_{avg})(Y - Y_{avg})}}{n-1} $$

Where `n` is the number of points in either X or Y.

#### Implement the following: 

<!-- grade_id: cell-31802707a6bea6cf, points: 0, grade: false, solution: true -->
```python
def covariance_top_half(s1, s2):
    """ 
    Top half of the fraction 
    
    Args:
        s1 (pd.Series): a pandas series 
        s2 (pd.Series): a pandas series (the same index and length as s1)

    Returns:
        top_half_value (float): value of the top half of the fraction 
    """    
    # Note: it is generally best to do this is multiple small steps 
    # e.g. start by making s1_avg, then later s1_minus_s1_avg, etc...
    
    # YOUR CODE HERE
    raise NotImplementedError()

def covariance_bottom_half(s1, s2):
    """ 
    Bottom half of the fraction 
    
    Args:
        s1 (pd.Series): a pandas series 
        s2 (pd.Series): a pandas series (the same index and length as s1)

    Returns:
        bottom_half_value (float): value of the bottom half of the fraction 
    """    
    # YOUR CODE HERE
    raise NotImplementedError()


def covariance_by_hand(s1, s2):
    top_half = covariance_top_half(s1, s2)
    bottom_half = covariance_bottom_half(s1, s2)
    return top_half / bottom_half
```

```python
a = pd.Series([1, 5, 7, 10, 25])
b = pd.Series([15, 30, 28, 45, 50])
print('Covariance by hand between a and b: %0.02f' % covariance_by_hand(a, b))
```

Expected output:   

    Covariance by hand between a and b: 112.05

<!-- grade_id: cell-110d6cc4cdab8695, points: 3, grade: true, solution: false -->
```python
a = pd.Series([1, 5, 7, 10, 25])
b = pd.Series([15, 30, 28, 45, 50]) 
c = pd.Series([22, 55, 23, 15, 92])
assert math.isclose(covariance_bottom_half(a, b), 4)
assert math.isclose(covariance_top_half(a, b), 448.2)
assert math.isclose(covariance_by_hand(a, b), 112.05)
assert math.isclose(covariance_by_hand(a, c), 230.2)
```

# Exercise 2: implement pearson correlation

Correlation is simply normalized covariance! 

$$ correlation = \frac{covariance(X, Y)}{\sqrt{Var(X) * Var(Y)}} $$

_Important note: if for some reason you really like to use numpy instead of pandas, please use ddof = 1. If this sentence didn't mean anything to you, just use pandas :)_ 

#### Complete here: 

<!-- grade_id: cell-7e2c47095ccc50ef, points: 0, grade: false, solution: true -->
```python
def pearson_correlation_by_hand(s1, s2): 
    """ 
    Naive implementation of pearson correlation
    
    Args:
        s1 (pd.Series): a pandas series 
        s2 (pd.Series): a pandas series (the same index and length as s1)

    Returns:
        pearson correlation (float): the correlation between s1 and s2 

    """
    
    # clue: remember, you've already implemented a function the top half in exercise one 
    # clue: for the square root you can use np.sqrt(...)
    
# YOUR CODE HERE
raise NotImplementedError()
   
    
    return pearson_correlation
```

```python
a = pd.Series([1, 5, 7, 10, 25])
b = pd.Series([15, 30, 28, 45, 50])
print('Correlation by hand between a and b: %0.02f' % pearson_correlation_by_hand(a, b))
```

Expected output:   

    Correlation by hand between a and b: 0.87

<!-- grade_id: cell-b7ab9bd4e5aa2dfa, points: 3, grade: true, solution: false -->
```python
assert round(pearson_correlation_by_hand(s1=pd.Series([2, 2, 7]), 
                                  s2=pd.Series([1, 12, 16])), 2) == .71
```

# Exercise 3: implement Spearman correlation 
Using your function `pearson_correlation_by_hand` calculate the Spearman correlation between `S1` and `S2`. 

So... what was Spearman correlation again?

_Note: don't use the pandas `.corr` to solve this, that would be cheating ;)_

<!-- grade_id: cell-c0748eb68560b07e, points: 0, grade: false, solution: true -->
```python
def spearman_correlation_using_pearson(s1, s2): 
    """ 
    Naive implementation of spearman correlation
    
    Args:
        s1 (pd.Series): a pandas series 
        s2 (pd.Series): a pandas series (the same index and length as s1)

    Returns:
        spearman_correlation (float): the correlation between s1 and s2 

    """
    
    # 
    # do the right transformation to s1 and s2
    # transformed_s1 = ...
    # transformed_s2 = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Use your pearson_correlation_by_hand function to get the spearman correlation 
    # spearman_correlation = ...
    # YOUR CODE HERE
    raise NotImplementedError()

    return spearman_correlation
```

```python
a = pd.Series([1, 5, 7, 10, 25])
b = pd.Series([15, 30, 28, 45, 50])
print('Spearman correlation between a and b: %0.02f' % 
      spearman_correlation_using_pearson(a, b))
```

Expected output:   

    Spearman correlation between a and b: 0.90

<!-- grade_id: cell-9c8a1bb04133d3d5, points: 3, grade: true, solution: false -->
```python
a = pd.Series([1, 5, 7, 10, 25])
b = pd.Series([15, 30, 28, 45, 50])
c = pd.Series([22, 55, 23, 15, 92])

assert math.isclose(spearman_correlation_using_pearson(a, b), a.corr(b, method='spearman'), abs_tol=.1)
assert math.isclose(spearman_correlation_using_pearson(b, c), b.corr(c, method='spearman'), abs_tol=.1)
assert math.isclose(spearman_correlation_using_pearson(c, a), c.corr(a, method='spearman'), abs_tol=.1)
```

----

# Exercise 4: enter an outlier 
We are going to analyze some house prices. In one case there will be no outliers, and in the other there will be a single outlier (maybe someone famous lived there) 

We will use regular pandas this time :) 

```python
plot_house_prices_and_rooms()
```

<!-- grade_id: cell-602c35acee50ad38, points: 0, grade: false, solution: true -->
```python
def calculate_correlations_with_pandas(house_prices, number_of_rooms):
    
    # pearson_corr = ...
    # YOUR CODE HERE
    raise NotImplementedError()

    # spearman_corr = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return pearson_corr, spearman_corr
```

```python
number_of_rooms, house_prices_normal, house_prices_with_outliers = get_house_prices_and_rooms()

pearson_corr_normal, spearman_corr_normal = calculate_correlations_with_pandas(house_prices_normal, 
                                                                               number_of_rooms)
pearson_corr_outlier, spearman_corr_outlier = calculate_correlations_with_pandas(house_prices_with_outliers, 
                                                                                 number_of_rooms)

# quick plot to see what happens
results = pd.Series({
    'Pearson without outlier': pearson_corr_normal,
    'Pearson with outlier': pearson_corr_outlier,
    'Spearman without outlier': spearman_corr_normal,
    'Spearman with outlier': spearman_corr_outlier,
})

results.plot(kind='barh')
plt.show()
```

Expected output 

    Pearson without outlier     0.689598
    Pearson with outlier        0.291450
    Spearman without outlier    0.635173
    Spearman with outlier       0.635284

<!-- grade_id: cell-1a7175bdcc21451a, points: 3, grade: true, solution: false -->
```python
assert math.isclose((pearson_corr_normal - pearson_corr_outlier), 0.398, abs_tol=.01)
assert math.isclose((spearman_corr_normal - spearman_corr_outlier), 0, abs_tol=.01)
```

# Exercise 5: Stocks! 

Time to predict the stockmarket and get ritch. 

<!-- grade_id: cell-de272afdcb6695e5, points: 0, grade: false, solution: true -->
```python
stocks = pd.read_csv('data/Stock_prices.csv')

# Explore the dataset. 
# Hint 1: you can use display(<name of df>) to force it to pretty print
# Hint 2: use the heatmap of the correlation matrix that we used in the learning notebooks 
# Hint 3: you may want to import something to help with the visualization 
# Hint 4: you can either paste the answers, or use a purely programmatic solution. 
# Hint 5: when we say lowest and highest we mean the sign, not the "magnitude"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-397d0b0c3105ffb6, points: 0, grade: false, solution: true -->
```python
# Complete the following questions 

# Q1: what is the pair of stocks with the highest pearson correlation? 
# (pass the answer as a list, and remember, you can just type it in, no fancy pandas needed) 
# stock_pair_with_highest_pearson_corr = ... 
# YOUR CODE HERE
raise NotImplementedError()

# # Q2: What is the stock pair with the lowest rank correlation with Occidental Petroleum? 
# (the answer is a string)
# lowest_rank_corr_with_occidental_petroleum = ... 
# YOUR CODE HERE
raise NotImplementedError()

# Q3: what is the pearson correlation between Marathon Oil Corp. and Occidental Petroleum? 
# (answer is a float)
# pearson_corr_between_marathon_oil_and_occidental_petr = ... 
# YOUR CODE HERE
raise NotImplementedError()

# Q4: Observe the top 4 pearson correlation pairs, and then look at the general correlation matrix. 
# Is there a confounding variable that may explain why these companies are so correlated? 
# possible_confounding_variable = 
# YOUR CODE HERE
raise NotImplementedError()
```

Expected output 

    No expected output on this one, see if it passes the test output ;) 

#### Test output (don't change code here) 

<!-- grade_id: cell-61ec8bb442300a8e, points: 3, grade: true, solution: false -->
```python
a = '9f94a681a8b473f6924ab5c42acaf8158fdb402c8d57a861f896dd17b09c2adb'
b = '52f85c820290f79a19a696f58cbb376e9034a21d66306565c29007facc7ab12d'
c = '0ab1e04f0684641e328c04b9485b4f50652e082ae4bb6541210e45c747a8f82f'

assert hash_answer(sorted(stock_pair_with_highest_pearson_corr)) == a
assert hash_answer(hash_answer(lowest_rank_corr_with_occidental_petroleum)) == b
assert math.isclose(pearson_corr_between_marathon_oil_and_occidental_petr, 0.9130, abs_tol=0.01)
assert hash_answer(possible_confounding_variable) == c
```

----

# Exercise 6: lots of stocks
You were hired by a hedge fund, because money. 

On the first day, your boss, Greedy Mc.Riskyface asked you to select one stock pair. He wants the two stocks to be as uncorrelated as possible (use pearson), so that he can short one and long the other. 

The answer should be (1) the two stocks, as a list and (2) their pearson correlation, as a float. 

<!-- grade_id: cell-ee4cfd1ff7abcd0e, points: 0, grade: false, solution: true -->
```python
stock_data = utils.get_stocks_data_2()

# most_uncorrelated_pair = ... 
# most_uncorrelated_pair_value = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6bc7867845b5cdfd, points: 5, grade: true, solution: false -->
```python
assert isinstance(most_uncorrelated_pair, list)
assert len(most_uncorrelated_pair) == 2
assert isinstance(most_uncorrelated_pair_value, float)
utils.dirty_little_secret()
```